|  CSE373 Analysis of Algorithm  |  Programming Practice #8 |  Posted on 11/26 Due 11/30   |
|:--|:--:|--:|

# Task #15

## Binomial Coefficients problem

We have covered Binomial Coefficients problem in the class, in this programming practice, you need to practice what we have learned in the class. Most of the codes are available on the class slide, you may use the code from the slide. This week’s practice is to play with the dynamic programming then you can familiarize with the DP coding.

**Solve Binomial Coefficients problem using <ins>Dynamic Programming (2D array)</ins>.**

#### Write your algorithm here:

<!-- grade_id: cell-f4b2b914f82e0838, points: 0, grade: false, solution: true -->
```python
def binomialCoeff_DP2D(n, k):
    """
    :type n: integer
    :type k: integer
    :rtype: integer
    """
    ### BEGIN SOLUTION
    poplist = [[-1 for x in range(k + 1)] for i in range(n + 1)]
    
    
    def recursion(n,k):
        if k == 0 or k == n:
            poplist[n][k] = 1
            return 1
        beforeC = poplist[n-1][k-1]
        afterC = poplist[n-1][k]
        if beforeC == -1:
            beforeC = recursion(n-1, k-1)
        if afterC == -1:
            afterC = recursion(n-1,k)
        poplist[n][k] = beforeC + afterC
        return beforeC + afterC
    
    return recursion(n, k)
    ### END SOLUTION
```

```python
n = 5
k = 2
print(binomialCoeff_DP2D(n , k))
```

Output:
```
10
```

#### Check the algorithm:

<!-- grade_id: cell-621215035cdfe527, points: 10, grade: true, solution: false -->
```python
n = [i*i for i in range(25)]
k = [i for i in range(25)]
answer = [1, 1, 6, 84, 1820, 53130, 1947792, 85900584, 4426165368, 260887834350, 17310309456440, 1276749965026536, 103619293824707388, 9176358300744339432, 880530516383349192480, 91005567811177478095440, 10078751602022313874633200, 1190739044344491048895397910, 149482492334195165714038760136, 19870867053543756004133247695400, 2788360983670896737872851072994080, 411887396336567398822620727355402190, 63887407766986865702182544710470036560, 10381758958529585222885358558747563185920, 1763783520005146433232953016554504214270600]
for (_n, _k, _answer) in zip(n, k, answer):
    i = binomialCoeff_DP2D(_n , _k)
    print("[Output] Value of C(%d,%d) is (%d)" %(_n, _k , i))
    assert i == _answer, "The result of your algorithm is %d where n=%d, k=%d. The right answer is %d" %(i, _n, _k, _answer)
```

Output:
```
[Output] Value of C(0,0) is (1)
[Output] Value of C(1,1) is (1)
[Output] Value of C(4,2) is (6)
[Output] Value of C(9,3) is (84)
[Output] Value of C(16,4) is (1820)
[Output] Value of C(25,5) is (53130)
[Output] Value of C(36,6) is (1947792)
[Output] Value of C(49,7) is (85900584)
[Output] Value of C(64,8) is (4426165368)
[Output] Value of C(81,9) is (260887834350)
[Output] Value of C(100,10) is (17310309456440)
[Output] Value of C(121,11) is (1276749965026536)
[Output] Value of C(144,12) is (103619293824707388)
[Output] Value of C(169,13) is (9176358300744339432)
[Output] Value of C(196,14) is (880530516383349192480)
[Output] Value of C(225,15) is (91005567811177478095440)
[Output] Value of C(256,16) is (10078751602022313874633200)
[Output] Value of C(289,17) is (1190739044344491048895397910)
[Output] Value of C(324,18) is (149482492334195165714038760136)
[Output] Value of C(361,19) is (19870867053543756004133247695400)
[Output] Value of C(400,20) is (2788360983670896737872851072994080
```

#### Check the execution time:

<!-- grade_id: cell-e34c8a9d417364b2, points: 1, grade: true, solution: false -->
```python
import time

start_time = time.time()
assert binomialCoeff_DP2D(1000, 500) == 270288240945436569515614693625975275496152008446548287007392875106625428705522193898612483924502370165362606085021546104802209750050679917549894219699518475423665484263751733356162464079737887344364574161119497604571044985756287880514600994219426752366915856603136862602484428109296905863799821216320

print("%s seconds" % (time.time() - start_time))
```

Output:
```
0.16674184799194336 seconds
```

```python

```