|  CSE373 Analysis of Algorithm  |  Programming Practice #14 |  Posted on 11/06 Due 11/11   |
|:--|:--:|--:|

# Task #14

## Binomial Coefficients problem

We have covered Binomial Coefficients problem in the class, in this programming practice, you need to practice what we have learned in the class. Most of the codes are available on the class slide, you may use the code from the slide. This week’s practice is to play with the dynamic programming then you can familiarize with the DP coding.

Solve Binomial Coefficients problem using _recursive function_.

#### Write your algorithm here:

<!-- grade_id: cell-03d95b8cf6bf25aa, points: 0, grade: false, solution: true -->
```python
def binomialCoeff_recursive(n, k):
    """
    :type n: integer
    :type k: integer
    :rtype: integer
    """
    ### BEGIN SOLUTION
    if k == 0 or k == n:
        return 1
    return binomialCoeff_recursive(n-1, k-1) + binomialCoeff_recursive(n-1,k)
    ### END SOLUTION
```

```python
n = 5
k = 2
print(binomialCoeff_recursive(n , k))
```

Output:
```
10
```

#### Check the algorithm:

<!-- grade_id: cell-799ded13b4dc5316, points: 10, grade: true, solution: false -->
```python
n = [1, 4, 9, 16, 25, 100, 10000]
k = [1, 2, 3, 4, 5, 0, 10000]
answer = [1, 6, 84, 1820, 53130, 1, 1]

for (_n, _k, _answer) in zip(n, k, answer):
    i = binomialCoeff_recursive(_n , _k)
    print("[Output] Value of C(%d,%d) is (%d)" %(_n, _k , i))
    assert i == _answer, "The result of your algorithm is %d where n=%d, k=%d. The right answer is %d" %(i, _n, _k, _answer)
```

Output:
```
[Output] Value of C(1,1) is (1)
[Output] Value of C(4,2) is (6)
[Output] Value of C(9,3) is (84)
[Output] Value of C(16,4) is (1820)
[Output] Value of C(25,5) is (53130)
[Output] Value of C(100,0) is (1)
[Output] Value of C(10000,10000) is (1)
```

#### Check the execution time:

<!-- grade_id: cell-40ae20dd77cbf73c, points: 5, grade: true, solution: false -->
```python
import time

start_time = time.time()

assert binomialCoeff_recursive(36, 7) == 8347680

print("%s seconds" % (time.time() - start_time))
```

Output:
```
2.2934021949768066 seconds
```

```python

```

```python

```