|  CSE373 Analysis of Algorithm  |  Programming Practice #13 |  Posted on 11/06 Due 11/11   |
|:--|:--:|--:|

# Task #13

## Solve Sudoku with backtracking

#### Guide:
1. Find all legal values of a given call
2. For each legal value, Go recursively and try to solve the grid

Your filled sudoku must be printable.

```python
grid=[[5,1,7,6,0,0,0,3,4],
        [2,8,9,0,0,4,0,0,0],
        [3,4,6,2,0,5,0,9,0],
        [6,0,2,0,0,0,0,1,0],
        [0,3,8,0,0,6,0,4,7],
        [0,0,0,0,0,0,0,0,0],
        [0,9,0,0,0,0,0,7,8],
        [7,0,3,4,0,0,5,6,0],
        [0,0,0,0,0,0,0,0,0]] # 9 x 9
```

#### Write your algorithm here:

<!-- grade_id: cell-077029b8f828c014, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

### END SOLUTION
def sudokuSolver(grid):
    #list1 = [0,0]
    temp = check_empty(grid)
    if (not temp):
        return True
    row = temp[0]
    col = temp[1]
    
    for num in range(1,10): #num 1~9
        if (is_it_safe(grid, row, col, num)):
            grid[row][col] = num
            if (sudokuSolver(grid)):
                return True
            grid[row][col] = 0
    return False
            
def check_empty(grid):
    for i in range(0, 9):
        for j in range(0, 9):
            if grid[i][j] == 0:
                #list1[0] = i #row
                #list1[1] = j #col
                #print("next empty:", (i, j))
                return [i, j]
    return False
def check_row(grid, row, num): # check whether it has the num already in the row
    #print("checking row: is", num, "in", grid[row])
    for i in range(0,9):
        if grid[row][i] == num:
            #print("yes")
            return True
    return False
def check_col(grid, col, num):
    #print("checking col:")
    for i in range(0,9):
        #print("is", grid[i][col], "equal to", num, ":", grid[i][col] == num)
        if grid[i][col] == num:
            return True
    return False
def check_box(grid, row, col, num):
    for i in range(row, row+3):
        for j in range(col, col+3):
            if grid[i][j] == num:
                return True
    return False
def is_it_safe(grid, row, col, num):
    return not check_row(grid, row, num) and not check_col(grid, col, num) and not check_box(grid, row - row%3, col - col%3, num)

```

```python
sudokuSolver(grid)
print(grid)

```

Output:
```
[[5, 1, 7, 6, 9, 8, 2, 3, 4], [2, 8, 9, 1, 3, 4, 7, 5, 6], [3, 4, 6, 2, 7, 5, 8, 9, 1], [6, 7, 2, 8, 4, 9, 3, 1, 5], [1, 3, 8, 5, 2, 6, 9, 4, 7], [9, 5, 4, 7, 1, 3, 6, 8, 2], [4, 9, 5, 3, 6, 2, 1, 7, 8], [7, 2, 3, 4, 8, 1, 5, 6, 9], [8, 6, 1, 9, 5, 7, 4, 2, 3]]
```

#### Check the algorithm:

```python
grid0_orig = [[0, 0, 0, 0, 0, 0, 6, 8, 0],
              [0, 0, 0, 0, 7, 3, 0, 0, 9],
              [3, 0, 9, 0, 0, 0, 0, 4, 5],
              [4, 9, 0, 0, 0, 0, 0, 0, 0],
              [8, 0, 3, 0, 5, 0, 9, 0, 2],
              [0, 0, 0, 0, 0, 0, 0, 3, 6],
              [9, 6, 0, 0, 0, 0, 3, 0, 8],
              [7, 0, 0, 6, 8, 0, 0, 0, 0],
              [0, 2, 8, 0, 0, 0, 0, 0, 0]]
grid1_orig = [[1, 0, 0, 2, 0, 0, 0, 4, 0],
              [0, 2, 0, 0, 0, 3, 9, 0, 0],
              [9, 0, 7, 0, 0, 0, 5, 0, 0],
              [0, 0, 4, 0, 0, 0, 0, 5, 7], 
              [0, 0, 0, 5, 4, 1, 0, 0, 0], 
              [3, 5, 0, 0, 0, 0, 1, 0, 0],
              [0, 0, 3, 0, 0, 0, 7, 0, 9], 
              [0, 0, 1, 4, 0, 0, 0, 8, 0], 
              [0, 9, 0, 0, 0, 2, 0, 0, 6]]
grid2_orig = [[5, 3, 0, 0, 7, 0, 0, 0, 0], 
              [6, 0, 0, 1, 9, 5, 0, 0, 0], 
              [0, 9, 8, 0, 0, 0, 0, 6, 0], 
              [8, 0, 0, 0, 6, 0, 0, 0, 3], 
              [4, 0, 0, 8, 0, 3, 0, 0, 1], 
              [7, 0, 0, 0, 2, 0, 0, 0, 6], 
              [0, 6, 0, 0, 0, 0, 2, 8, 0], 
              [0, 0, 0, 4, 1, 9, 0, 0, 5], 
              [0, 0, 0, 0, 8, 0, 0, 7, 9]]
grid3_orig = [[0, 0, 0, 0, 0, 0, 0, 0, 0], 
              [0, 9, 3, 6, 2, 8, 1, 4, 0], 
              [0, 6, 0, 0, 0, 0, 0, 5, 0], 
              [0, 3, 0, 0, 1, 0, 0, 9, 0], 
              [0, 5, 0, 8, 0, 2, 0, 7, 0], 
              [0, 4, 0, 0, 7, 0, 0, 6, 0], 
              [0, 8, 0, 0, 0, 0, 0, 3, 0],
              [0, 1, 7, 5, 9, 3, 4, 2, 0], 
              [0, 0, 0, 0, 0, 0, 0, 0, 0]]
```

<!-- grade_id: cell-d7fa65fe64db7ee2, points: 10, grade: true, solution: false -->
```python
from copy import deepcopy

grid0 = deepcopy(grid0_orig)
grid1 = deepcopy(grid1_orig)
grid2 = deepcopy(grid2_orig)
grid3 = deepcopy(grid3_orig)

sudokuSolver(grid0)
sudokuSolver(grid1)
sudokuSolver(grid2)
sudokuSolver(grid3)

assert grid0 == [[1, 7, 2, 5, 4, 9, 6, 8, 3],
                 [6, 4, 5, 8, 7, 3, 2, 1, 9],
                 [3, 8, 9, 2, 6, 1, 7, 4, 5],
                 [4, 9, 6, 3, 2, 7, 8, 5, 1],
                 [8, 1, 3, 4, 5, 6, 9, 7, 2],
                 [2, 5, 7, 1, 9, 8, 4, 3, 6],
                 [9, 6, 4, 7, 1, 5, 3, 2, 8],
                 [7, 3, 1, 6, 8, 2, 5, 9, 4],
                 [5, 2, 8, 9, 3, 4, 1, 6, 7]]
assert grid1 == [[1, 3, 5, 2, 9, 7, 6, 4, 8], 
                 [4, 2, 6, 8, 5, 3, 9, 7, 1], 
                 [9, 8, 7, 6, 1, 4, 5, 3, 2], 
                 [8, 1, 4, 9, 3, 6, 2, 5, 7], 
                 [7, 6, 2, 5, 4, 1, 8, 9, 3], 
                 [3, 5, 9, 7, 2, 8, 1, 6, 4], 
                 [6, 4, 3, 1, 8, 5, 7, 2, 9], 
                 [2, 7, 1, 4, 6, 9, 3, 8, 5], 
                 [5, 9, 8, 3, 7, 2, 4, 1, 6]]
assert grid2 == [[5, 3, 4, 6, 7, 8, 9, 1, 2], 
                 [6, 7, 2, 1, 9, 5, 3, 4, 8], 
                 [1, 9, 8, 3, 4, 2, 5, 6, 7], 
                 [8, 5, 9, 7, 6, 1, 4, 2, 3], 
                 [4, 2, 6, 8, 5, 3, 7, 9, 1], 
                 [7, 1, 3, 9, 2, 4, 8, 5, 6], 
                 [9, 6, 1, 5, 3, 7, 2, 8, 4], 
                 [2, 8, 7, 4, 1, 9, 6, 3, 5], 
                 [3, 4, 5, 2, 8, 6, 1, 7, 9]]
assert grid3 == [[2, 7, 1, 9, 5, 4, 6, 8, 3], 
                 [5, 9, 3, 6, 2, 8, 1, 4, 7], 
                 [4, 6, 8, 1, 3, 7, 2, 5, 9], 
                 [7, 3, 6, 4, 1, 5, 8, 9, 2], 
                 [1, 5, 9, 8, 6, 2, 3, 7, 4], 
                 [8, 4, 2, 3, 7, 9, 5, 6, 1], 
                 [9, 8, 5, 2, 4, 1, 7, 3, 6], 
                 [6, 1, 7, 5, 9, 3, 4, 2, 8], 
                 [3, 2, 4, 7, 8, 6, 9, 1, 5]]
```

#### Check the execution time:

```python
grid0 = deepcopy(grid0_orig)
grid1 = deepcopy(grid1_orig)
grid2 = deepcopy(grid2_orig)
grid3 = deepcopy(grid3_orig)

import time

start_time = time.time()

sudokuSolver(grid0)
sudokuSolver(grid1)
sudokuSolver(grid2)
sudokuSolver(grid3)

print("%s seconds" % (time.time() - start_time))
```

Output:
```
0.4024677276611328 seconds
```

```python

```

```python

```